// variables4.rs
// Make me compile! Execute the command `rustlings hint variables4` if you want a hint :)

// ✓ I AM NOT DONE

fn main() {
    //let x: i32;
    let x: i32 = 5;
    println!("Number {}", x);
}

// $ rustlings hint variables4
/*
Oops! In this exercise, we have a variable binding that we've created on   
line 7, and we're trying to use it on line 8, but we haven't given it a    
value. We can't print out something that isn't there; try giving x a value!
This is an error that can cause bugs that's very easy to make in any       
programming language -- thankfully the Rust compiler has caught this for us!
*/