// variables1.rs
// Make me compile!
// Execute the command `rustlings hint variables1` if you want a hint :)

// ✓ I AM NOT DONE

fn main() {
    //x = 5;
    let x = 5;
    println!("x has the value {}", x);
}

// $ rustlings hint variables1
/*
Hint: The declaration on line 12 is missing a keyword that is needed in Rust
to create a new variable binding.
*/