// quiz2.rs
// This is a quiz for the following sections:
// - Strings

// Ok, here are a bunch of values-- some are `String`s, some are `&str`s. Your
// task is to call one of these two functions on each value depending on what
// you think each value is. That is, add either `string_slice` or `string`
// before the parentheses on each line. If you're right, it will compile!

// ✓ I AM NOT DONE

fn string_slice(arg: &str) {
    println!("{}", arg);
}
fn string(arg: String) {
    println!("{}", arg);
}

fn main() {
    //???("blue");
    string_slice("blue");
    //???("red".to_string());
    string("red".to_string());
    //???(String::from("hi"));
    string(String::from("hi"));
    //???("rust is fun!".to_owned());
    string("rust is fun!".to_owned());
    //???("nice weather".into());
    string("nice weather".into());
    //???(format!("Interpolation {}", "Station"));
    string(format!("Interpolation {}", "Station"));
    //???(&String::from("abc")[0..1]);
    string_slice(&String::from("abc")[0..1]);
    //???("  hello there ".trim());
    string_slice("  hello there ".trim());
    //???("Happy Monday!".to_string().replace("Mon", "Tues"));
    string("Happy Monday!".to_string().replace("Mon", "Tues"));
    //???("mY sHiFt KeY iS sTiCkY".to_lowercase());
    string("mY sHiFt KeY iS sTiCkY".to_lowercase());
}

// $ rustlings hint quiz2
// No hints this time ;)