// functions1.rs
// Make me compile! Execute `rustlings hint functions1` for hints :)

// ✓ I AM NOT DONE

fn call_me() {}

fn main() {
    //let call_me = || {};
    call_me();
}

// $ rustlings hint functions1
/*
This main function is calling a function that it expects to exist, but the  
function doesn't exist. It expects this function to have the name `call_me`.
It expects this function to not take any arguments and not return a value.  
Sounds a lot like `main`, doesn't it?
*/